```python
# Initialize Otter
import otter
grader = otter.Notebook("hw08_humboldt.ipynb")
```

# Homework 8: Confidence Intervals

**Helpful Resource:**

- [Python Reference](https://www.data8.org/datascience/reference-nb/datascience-reference.html): Cheat sheet of helpful array & table methods used in Data 111!

**Recommended Reading**: 

* [Estimation](https://www.inferentialthinking.com/chapters/13/Estimation)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the instructor or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *
import d8error

# These lines do some fancy plotting magic.",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1. Italian Restaurants in Arcarta/Eureka

Bonnie and Bethany are trying to see what the best Italian restaurant is in the area. They survey 1,500 Cal Poly Humboldt students selected uniformly at random and ask each student what Italian restaurant is the best. (*Note: This data is fabricated for the purposes of this homework.*) The choices of Italian restaurants are [Mazzotti's Italian Restaurant](https://www.google.com/maps/place/Mazzotti's+-+Italian+Restaurant/@40.8681176,-124.0868234,15z/data=!4m6!3m5!1s0x54d1570f63be6e0b:0x55e3986e9f3226f4!8m2!3d40.8681176!4d-124.0868234!16s%2Fg%2F1v27v6zr?entry=ttu), [Brick & Fire](https://www.google.com/maps/place/Brick+%26+Fire/@40.7923298,-124.166398,17z/data=!3m1!4b1!4m6!3m5!1s0x54d3ffddc91d64df:0x8a806dade51653e9!8m2!3d40.7923258!4d-124.1638231!16s%2Fg%2F11f55d5rg7?entry=ttu), [Babe's Pizza & Pasta](https://www.google.com/maps/place/Babe's+Pizza+%26+Pasta/@40.7694284,-124.145214,17z/data=!3m1!4b1!4m6!3m5!1s0x54d3ff9e61a3a4ad:0xb708b698daf884b!8m2!3d40.7694244!4d-124.1426391!16s%2Fg%2F1tf6lyln?entry=ttu), and [Gabriel's Italian Restaurant](https://www.google.com/maps/place/Gabriel%E2%80%99s+Italian+Restaurant/@40.8038929,-124.1703527,17z/data=!3m1!4b1!4m6!3m5!1s0x54d1556528cd9149:0xf35a84f0bb853f90!8m2!3d40.8038889!4d-124.1677778!16s%2Fg%2F1tqg01jw?entry=ttu). After compiling the results, Bonnie and Bethany release the following percentages from their sample:

|Italian Restaurant  | Percentage|
|:------------:|:------------:|
|Mazzotti's | 8% |
|Brick & Fire | 53% |
|Babe's Pasta & Pizza | 25% |
|Gabriel's | 14% |

These percentages represent a uniform random sample of the population of Cal Poly Humboldt students. We will attempt to estimate the corresponding *parameters*, or the percentage of the votes that each restaurant would receive from the population (i.e. all Cal Poly Humboldt students). We will use confidence intervals to compute a range of values that reflects the uncertainty of our estimates.

The table `votes` contains the results of Bonnie and Bethany's survey.

```python
# Just run this cell
votes = Table.read_table('votes.csv')
votes
```

```python
# Just run this cell
votes = Table.read_table('votes.csv')
votes
```

**Question 1.1.** Complete the function `one_resampled_percentage` below. It should return Brick & Fire's ***percentage*** of votes after taking the original table (`tbl`) and performing one bootstrap sample of it. Reminder that a percentage is between 0 and 100. **(9 Points)**

*Note:* `tbl` will always be in the same format as `votes`.

*Hint:* Given a table of votes, how can you figure out what percentage of the votes are for a certain restaurant? **Be sure to use percentages, not proportions, for this question!**


```python
def one_resampled_percentage(tbl):
    ...

one_resampled_percentage(votes)
```

```python
grader.check("q1_1")
```

**Question 1.2.** Complete the `percentages_in_resamples` function such that it simulates and returns an array of 2022 elements, where each element represents a bootstrapped estimate of the percentage of voters who will vote for Brick & Fire. You should use the `one_resampled_percentage` function you wrote above. **(9 Points)**


```python
def percentages_in_resamples():
    percentage_brick_fire = make_array()
    ...
```

```python
grader.check("q1_2")
```

In the following cell, we run the function you just defined, `percentages_in_resamples`, and create a histogram of the calculated statistic for the 2022 bootstrap estimates of the percentage of voters who voted for Brick & Fire. 

*Note:* This might take a few seconds to run.

```python
resampled_percentages = percentages_in_resamples()
Table().with_column('Estimated Percentage', resampled_percentages).hist("Estimated Percentage")
```

**Question 1.3.** Using the array `resampled_percentages`, find the values at the two edges of the middle 95% of the bootstrapped percentage estimates. Compute the lower and upper ends of the interval, named `brick_fire_lower_bound` and `brick_fire_upper_bound` respectively. **(9 Points)**

*Hint:* If you are stuck on this question, try looking over [Chapter 13](https://inferentialthinking.com/chapters/13/Estimation.html) of the textbook.


```python
brick_fire_lower_bound = ...
brick_fire_upper_bound = ...
print(f"Bootstrapped 95% confidence interval for the percentage of Brick & Fire voters in the population: [{brick_fire_lower_bound:.2f}, {brick_fire_upper_bound:.2f}]")
```

```python
grader.check("q1_3")
```

**Question 1.4.** The survey results seem to indicate that Brick & Fire is beating all the other Italian restaurants among the voters. We would like to use confidence intervals to determine a range of likely values for Brick & Fire's true lead over all the other restaurants combined. The calculation for Brick & Fire's lead over Mazzotti's, Babe's, and Gabriel's combined is:

$$\text{Brick & Fire's percent of the vote} - (\text{100 percent} - \text{Brick & Fire's percent of vote})$$

Define the function `one_resampled_difference` that returns **exactly one value** of Brick & Fire's percentage lead over Mazzotti's, Babe's, and Gabriel's combined from one bootstrap sample of `tbl`. **(9 Points)**

*Hint 1:* Brick & Fire's lead can be negative.

*Hint 2:* Given a table of votes, how can you figure out what percentage of the votes are for a certain restaurant? **Be sure to use percentages, not proportions, for this question!**

*Note:* If the skeleton code provided within the function is not helpful for you, feel free to approach the question using your own variables.


```python
def one_resampled_difference(tbl):
    bootstrap = ...
    brick_fire_percentage = ...
    ...
```

```python
grader.check("q1_4")
```

<!-- BEGIN QUESTION -->

**Question 1.5.** Write a function called `leads_in_resamples` that returns an array of 2022 elements representing the bootstrapped estimates (the result of calling `one_resampled_difference`) of Brick & Fire's lead over Mazzotti's, Babe's, and Gabriel's combined. Afterwards, run the cell to plot a histogram of the resulting samples. **(9 Points)**

*Hint:* If you see an error involving `NoneType`, consider what components a function needs to have!


```python
def leads_in_resamples():
    ...

sampled_leads = leads_in_resamples()
Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")
```

<!-- END QUESTION -->

**Question 1.6.** Use the simulated data in `sampled_leads` from Question 1.5 to compute an approximate 95% confidence interval for Brick & Fire's true lead over Mazzotti's, Babe's, and Gabriel's combined. **(9 Points)**


```python
diff_lower_bound = ...
diff_upper_bound = ...
print("Bootstrapped 95% confidence interval for Brick & Fire's true lead over Mazzotti's, Babe's, and Gabriel's combined: [{:f}%, {:f}%]".format(diff_lower_bound, diff_upper_bound))
```

```python
grader.check("q1_6")
```

## 2. Interpreting Confidence Intervals 

Dr. Johnson computed the following 95% confidence interval for the percentage of Brick & Fire voters: 

$$[50.53, 55.53]$$

(Your answer may have been a bit different due to randomness; that doesn't mean it was wrong!)

<!-- BEGIN QUESTION -->

**Question 2.1.** Dr. Johnson also created 70%, 90%, and 99% confidence intervals from the same sample, but we forgot to label which confidence interval represented which percentages! **First**, match each confidence level (70%, 90%, 99%) with its corresponding interval in the cell below (e.g. __ % CI: [52.1, 54] $\rightarrow$ replace the blank with one of the three confidence levels). **Then**, explain your thought process and how you came up with your answers. **(10 Points)**

The intervals are below:

* [50.03, 55.94]
* [52.1, 54]
* [50.97, 54.99]



_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.2.** Suppose we produced 6,000 new samples (each one a new/distinct uniform random sample of 1,500 students) from the population and created a 95% confidence interval from each one. Roughly how many of those 6,000 intervals do you expect will actually contain the true percentage of the population? **(9 Points)**

Assign your answer to `true_percentage_intervals`.


```python
true_percentage_intervals = ...
```

```python
grader.check("q2_2")
```

Recall the second bootstrap confidence interval you created, which estimated Brick & Fire's lead over Mazzotti's, Babe's, and Gabriel's combined. Among
voters in the sample, Brick & Fire's lead was 6%. Dr. Johnson's 95% confidence interval for the true lead (in the population of all voters) was:

$$[1.2, 11.2]$$

Suppose we are interested in testing a simple yes-or-no question:

> "Is the percentage of votes for Brick & Fire equal to the percentage of votes for Mazzotti's, Babe's, and Gabriel's combined?"

Our null hypothesis is that the percentages are equal, or equivalently, that Brick & Fire's lead is exactly 0. Our alternative hypothesis is that Brick & Fire's lead is not equal to 0.  In the questions below, don't compute any confidence interval yourself—use only Dr. Johnson's 95% confidence interval.

**Question 2.3.** Say we use a 5% p-value cutoff. Do we reject the null, fail to reject the null, or are we unable to tell using the Dr. Johnson's confidence interval? **(9 Points)**

Assign `cutoff_five_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using Dr. Johnson's confidence interval

*Hint:* Consider the relationship between the p-value cutoff and confidence. If you're confused, take a look at [this chapter](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) of the textbook.


```python
cutoff_five_percent = ...
```

```python
grader.check("q2_3")
```

**Question 2.4.** What if, instead, we use a p-value cutoff of 1%? Do we reject the null, fail to reject the null, or are we unable to tell using the confidence interval? **(9 Points)**

Assign `cutoff_one_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using Dr. Johnson's confidence interval


```python
cutoff_one_percent = ...
```

```python
grader.check("q2_4")
```

**Question 2.5.** What if we use a p-value cutoff of 10%? Do we reject, fail to reject, or are we unable to tell using the confidence interval? **(9 Points)**

Assign `cutoff_ten_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using Dr. Johnson's confidence interval


```python
cutoff_ten_percent = ...
```

```python
grader.check("q2_5")
```

You're done with Homework 8!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment on Canvas. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0, 2])
```
>>> type(one_resampled_percentage(votes)) in set([float, np.float64]) 
True
```
```
>>> # Remember, the question asks for a percentage, not a proportion. 
>>> one_resampled_percentage(votes) <= 0
False
```
```
>>> 35 <= one_resampled_percentage(votes) <= 65
True
```


Tests for q1_2 (points: [0])
```
>>> len(percentages_in_resamples()) == 2022
True
```


Tests for q1_3 (points: [0])
```
>>> 45 <= brick_fire_lower_bound <= brick_fire_upper_bound <= 65
True
```


Tests for q1_4 (points: [0])
```
>>> type(one_resampled_difference(votes)) in set([float, np.float64]) 
True
```


Tests for q1_6 (points: [0])
```
>>> -5 <= diff_lower_bound <= diff_upper_bound <= 20
True
```


Tests for q2_2 (points: [0])
```
>>> 1000 <= true_percentage_intervals <= 10000
True
```


Tests for q2_3 (points: [0, 0, 0])
```
>>> 1 <= cutoff_five_percent <= 3
True
```
```
>>> # Remember, we are using a 95% confidence interval of [1.2, 11.2].
>>> # Our null hypothesis claims that Brick & Fire's lead is 0.
>>> # This falls outside of our 95% confidence interval.
>>> # What can we conclude if we use a 5% p-value cutoff? 
>>> cutoff_five_percent == 3
False
```
```
>>> # We are using a 95% confidence interval of [1.2, 11.2].
>>> # Our null hypothesis claims that Brick & Fire's lead is 0.
>>> # This falls outside of our 95% confidence interval. 
>>> # 100% - 95% = 5%, which is our p-value cutoff.
>>> # At the 5% level of significance, 0 doesn't seem like a plausible value
>>> # for Brick & Fire's lead. We would reject the null.
>>> # Remember, there is a duality between confidence intervals and tests:
>>> # if you are testing whether or not the true lead is a particular value x,
>>> # and you use the 5% cutoff for the p-value, then you will reject the null
>>> # if x is not in your 95% confidence interval for the lead.
>>> cutoff_five_percent == 2
False
```


Tests for q2_4 (points: [0])
```
>>> 1 <= cutoff_one_percent <= 3
True
```


Tests for q2_5 (points: [0])
```
>>> 1 <= cutoff_ten_percent <= 3
True
```
