```python
# Initialize Otter
import otter
grader = otter.Notebook("hw10_humboldt.ipynb")
```

# Homework 10: Linear Regression

**Helpful Resource:**

- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in Data 8!

**Recommended Readings**: 

* [The Regression Line](https://www.inferentialthinking.com/chapters/15/2/Regression_Line.html)
* [Method of Least Squares](https://www.inferentialthinking.com/chapters/15/3/Method_of_Least_Squares.html)
* [Least Squares Regression](https://www.inferentialthinking.com/chapters/15/4/Least_Squares_Regression.html)

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to setup the notebook by importing some helpful libraries. Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. **Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Note: This homework has hidden tests on it. That means even though the tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.**


Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

```python
# Run this cell to set up the notebook, but please don't change it.

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
from datetime import datetime
```

## 1. Triple Jump Distances vs. Vertical Jump Heights 

Does skill in one sport imply skill in a related sport?  The answer might be different for different activities. Let's find out whether it's true for the [triple jump](https://en.wikipedia.org/wiki/Triple_jump) (a horizontal jump similar to a long jump) and the [vertical jump](https://en.wikipedia.org/wiki/Vertical_jump).  Since we're learning about linear regression, we will look specifically for a *linear* association between skill level in the two sports.

The following data was collected by observing 40 collegiate-level soccer players. Each athlete's distances in both events were measured in centimeters. Run the cell below to load the data.

```python
# Run this cell to load the data
jumps = Table.read_table('triple_vertical.csv')
jumps
```

**Question 1.1.** Create a function `standard_units` that converts the values in the array `data` to standard units. **(5 points)**


```python
def standard_units(data):
    ...
```

```python
grader.check("q1_1")
```

**Question 1.2.** Now, using the `standard_units` function, define the function `correlation` which computes the correlation between `x` and `y`. **(5 points)**


```python
def correlation(x, y):
    ...
```

```python
grader.check("q1_2")
```

<!-- BEGIN QUESTION -->

**Question 1.3.** Before running a regression, it's important to see what the data looks like, because our eyes are good at picking out unusual patterns in data.  Draw a scatter plot, **that includes the regression line**, with the triple jump distances on the horizontal axis and the vertical jump heights on vertical axis. **(5 points)**

See the documentation on `scatter` [here](http://data8.org/datascience/_autosummary/datascience.tables.Table.scatter.html#datascience.tables.Table.scatter) for instructions on how to have Python draw the regression line automatically.

*Hint:* The `fit_line` argument may be useful here!


```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 1.4.** Does the correlation coefficient $r$ look closest to 0, .5, or -.5? Explain. **(5 points)**


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 1.5.** Create a function called `parameter_estimates` that takes in the argument `tbl`, a two-column table where the first column is the x-axis and the second column is the y-axis. It should return an array with three elements: the **(1) correlation coefficient** of the two columns and the **(2) slope** and **(3) intercept** of the regression line that predicts the second column from the first, in original units. **(5 points)**

*Hint:* This is a rare occasion where it’s better to implement the function using column indices instead of column names, in order to be able to call this function on any table. If you need a reminder about how to use column indices to pull out individual columns, please refer to [this](https://www.inferentialthinking.com/chapters/06/Tables.html#accessing-the-data-in-a-column) section of the textbook.


```python
def parameter_estimates(tbl):
    ...
    return make_array(r, slope, intercept)

parameters = parameter_estimates(jumps) 
print('r:', parameters.item(0), '; slope:', parameters.item(1), '; intercept:', parameters.item(2))
```

```python
grader.check("q1_5")
```

**Question 1.6.** Now suppose you want to go the other way and predict a triple jump distance given a vertical jump distance. What would the regression parameters of this linear model be? How do they compare to the regression parameters from the model where you were predicting vertical jump distance given a triple jump distance (in Question 1.5)? **(5 points)**

Set `regression_changes` to an array of 3 elements, with each element corresponding to whether or not the corresponding item returned by `parameter_estimates` changes when switching vertical and triple as $x$ and $y$. For example, if $r$ changes, the slope changes, but the intercept wouldn't change, the `regression_changes` would be assigned to `make_array(True, True, False)`.


```python
regression_changes = ...
regression_changes
```

```python
grader.check("q1_6")
```

**Question 1.7.** Let's use `parameters` (from Question 1.5) to predict what certain athletes' vertical jump heights would be given their triple jump distances. **(5 points)**

The world record for the triple jump distance is 18.29 *meters* by Johnathan Edwards. What is the prediction for Edwards' vertical jump using this line?

*Hint:* Make sure to convert from meters to centimeters!


```python
triple_record_vert_est = ...
print("Predicted vertical jump distance: {:f} centimeters".format(triple_record_vert_est))
```

```python
grader.check("q1_7")
```

<!-- BEGIN QUESTION -->

**Question 1.8.** Do you think it makes sense to use this line to predict Edwards' vertical jump? **(5 points)**

*Hint:* Compare Edwards' triple jump distance to the triple jump distances in `jumps`. Is it relatively similar to the rest of the data (shown in Question 1.3)? 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Relative Sea Level Rise

#### Global Mean Sea Level

Humans are altering Earth's climate by increasing the amount of heat-trapping greenhouse gases in the atmosphere. Atmospheric carbon dioxide (the most commonly emitted greenhouse gas) has reached 400 parts per million, a level unseen on earth in the last 800,000 years. Most of the additional heat that is being trapped is being absorbed by the oceans. As the water warms up, it expands, takes up more space, and causes sea levels to rise. Additionally, melting land ice (especially on Greenland and Antarctica) is adding mass to the ocean, further raising **global mean sea level**(GMSL).

<!-- BEGIN QUESTION -->

**Question 2.1.** Based on the description above, what are the two main ways that climate change increases global mean sea level? **(2 points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

#### Relative Sea Level

Each coastal location will experience sea level rise differently due to local and regional factors. One of these factors is the rising or sinking of land due to plate tectonics, called **vertical land motion** (VLM). If the land is sinking at a coastal location AND the ocean in the region is rising, then the locals might be very worried. If instead the land is rising faster than the ocean is rising, then the locals might feel as if the ocean is receeding. We can account for this effect by using **relative sea level** (RSL), where we observe water levels relative to a land based reference frame. 

#### Relative Sea Level in Northern California

Let's take a look at relative sea level data from two different tide gauge locations along our coast to see how this plays out...

```python
stations = Table.read_table("NorCalSLRmeta.csv")

Marker.map_table(stations.select('lat', 'lon', 'station_name').relabel('station_name', 'labels'))
```

Zoom out on the map until you see both blue markers. Click on each marker to see the labels for the tide gauge stations. 

<!-- BEGIN QUESTION -->

**Question 2.2.** What is the name of the tide gauge station that is closest to Cal Poly Humboldt? **(1 point)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Now let's look at the data from these stations. The data is in meters and is an average of the water levels for each month. Additionally, regular seasonal fluctuations have been removed.

```python
rsl = Table.read_table("NorCalSLR.csv")
rsl
```

Here is a plot of both stations:

```python
rsl.plot("months since December 1979",["North Spit", "Crescent City"])
plt.title('Monthly Avg RSL (Seasonal Cycle Removed) ')
plt.ylabel('RSL [meters]')
```

Let's take a closer look at each of these stations, one by one.

## 3. Sea Level Rise at North Spit

**Question 3.1.** Use the function `parameter_estimates` that you defined earlier to compute the least-squares linear regression line that predicts relative sea level (RSL) at *North Spit* based on the month, in original units. We have provided a two column table for you in the cell below with the first column representing the months (x) and the second column representing observed RSL at North Spit (y), which you should use as an argument to the function. Use the output from the `parameter_estimates` function to find the slope (`rsl_ns_slope`) and intercept (`rsl_ns_intercept`) of the regression line. **(2 points)**

```python
rsl_ns_tbl = rsl.select('months since December 1979', 'North Spit').relabeled('North Spit','Observed')
estimates = ...
rsl_ns_slope = ...
rsl_ns_intercept = ...
print("Slope:", round(rsl_ns_slope, 5))
print("Intercept", round(rsl_ns_intercept, 5))
```

```python
grader.check("q3_1")
```

**Question 3.2.** Convert the slope for North Spit from *meters per month* **to** *millimeters per year* (which are the units typically used to document sea level rise) (**1 point**)

```python
rsl_ns_slope_mm_per_year = ...
print("Sea Level Rise [mm/year] at North Spit:", round(rsl_ns_slope_mm_per_year,2))
```

This number is the *relative sea level **rise*** estimate for *North Spit*. Wigi, as named by the Wiyot people, is also known as Humboldt Bay. As you can see from the map we made above, the North Spit tide gauge is located in Wigi. **Wigi is experiencing one of the fastest rates of relative sea level rise on the West Coast!** This is mostly due to the sinking (subsidence) of the tectonic plate here!

**Question 3.3.** Write a function `linear_model` which has three inputs: a dependent variable array (x), slope value, and intercept value. It should output an array of independent values (y) for each dependent value. (**2 points**)

```python
def linear_model(x,slope,intercept):
    ...
```

```python
grader.check("q3_3")
```

**Question 3.4.** Write a function `linear_fit` which takes the table that consists of an independent variable and a dependent variable as input (for example your `rsl_ns_tbl` should work). The function should use the first column of the table to predict the second column of the table using the linear regression model. For `rsl_ns_tbl` for example, it should output a prediction that is an array of numbers associated with each month in the table. Again, it will be helpful to use the functions `parameter_estimates` and `linear_model` that you defined earlier in this homework. **(5 points)**

*Note:* Make sure that your `linear_fit` function is using least squares linear regression.

```python
def linear_fit(tbl):
    parameters = ...
    slope = ...
    intercept = ...
    ...
```

```python
grader.check("q3_4")
```

<!-- BEGIN QUESTION -->

**Question 3.5.** Now, using the `linear_fit` function you just defined, make a plot with "months since December 1979" along the x-axis and both real and modeled relative sea level along the y-axis. The color of the lines for the real sea level should be different from the color for the predicted sea level. Make sure to label your axes and note your units. (5 points)

*Hint 1*: An example of a line plot can be found above.

*Hint 2*: Think about the table that must be used to generate this line plot. What data should the columns represent? Based on the data that you need, how many columns should be present in this table? Also, what should each row represent? Constructing the table will be the main part of this question; once you have this table, generating the line plot should be more straightforward.

```python
# BEGIN SOLUION
rsl_ns_tbl = rsl_ns_tbl.with_column("Linear Fit", linear_fit(rsl_ns_tbl)) 
rsl_ns_tbl.plot('months since December 1979',["Observed","Linear Fit"]) 
plt.title('Monthly Average RSL at North Spit (Seasonal Removed)')
plt.ylabel('RSL [meters]')
#END SOLUTION
```

<!-- END QUESTION -->

**Question 3.6.** Add another column to `rsl_ns_tbl` called "Fit Residual" with values of the sea level Observations minus the Linear Fit prdictions at North Spit. We will look at this residual in a later section of the activity. **(1 point)**

```python
rsl_ns_tbl = ...
rsl_ns_tbl
```

**Question 3.7.** Next create another function called `linear_predictor` that can predict values outside of the domain of the observations. (For example, what if we wanted to know what the sea levels will be in 2100?) This function will be very similar to the linear_fit function you just created but the input and output will change. You will still need to input a table consisting of independent and dependent values, but you will need one additional input representing the new independent value where you want a prediction. There will still be one output variable but it will provide linearly predicted values for the new indepedent values that were input. **(5 points)**

```python
def linear_predictor(tbl,x):
    parameters = ...
    slope = ...
    intercept = ...
    ...
```

```python
grader.check("q3_7")
```

**Question 3.8.** Now use your `linear_predictor` function to estimate relative sea level at North Spit in meters for the year 2100 (which is 1453 months after December 1979.) **(1 point)**

```python
...
```

Here are some official estimates of how sea level is projected to change in the future at *North Spit*, where each color represents a different green house gas emissions scenario. *Note that sea level is shown in meters on the right*

<img src="NorthSpitSLRprojection.jpg" alt="drawing" width="1000"/>

You can find more info on these projections [here](https://oceanservice.noaa.gov/hazards/sealevelrise/sealevelrise-tech-report-sections.html).

<!-- BEGIN QUESTION -->

**Question 3.9.** Based on your estimate for 2100 and also the shape of these prediction curves, do you think your linear model is reasonable for North Spit? Why or why not? **(5 points)**

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Sea Level Rise at Crescent City

Now let's do something similar for *Crescent City*. First let's make a new table that just has the Crescent City data alongside each month.

```python
rsl_cc_tbl = rsl.select('months since December 1979', 'Crescent City').relabeled('Crescent City','Observed')
rsl_cc_tbl
```

Again we will use linear regression to estimate sea level rise here. However, this time we will program our linear regression a slightly different way. Recall that the least-squares regression line is the unique straight line that minimizes root mean squared error (RMSE) among all possible fit lines. Using this property, we can find the equation of the regression line by finding the pair of slope and intercept values that minimize root mean squared error. 

You're done with Homework 10!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment according to your instructor's directions. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0, 0])
```
>>> standard_units([1,2,3,4,5])
array([-1.41421356, -0.70710678,  0.        ,  0.70710678,  1.41421356])
```
```
>>> np.isclose(np.mean(standard_units([1,2,3,4,5])), 0)
True
```
```
>>> np.isclose(np.std(standard_units([1,2,3,4,5])), 0.9999999999999999)
True
```


Tests for q1_2 (points: [0])
```
>>> np.isclose(correlation([1,2,3], [4,5,6]), 0.9999999999999999)
True
```


Tests for q1_5 (points: [0, 0, 0])
```
>>> len(parameters) == 3
True
```
```
>>> # Make sure your function is returning an array!
>>> type(parameters) == np.ndarray
True
```
```
>>> np.isclose(parameters.item(0), 0.8343076972837598)
True
```


Tests for q1_6 (points: [0, 0])
```
>>> # Make sure regression_changes is a numpy array with 3 items!
>>> len(regression_changes) == 3 and type(regression_changes) == np.ndarray
True
```
```
>>> regression_changes.item(0) in set([0, 1, False, True])
True
```


Tests for q1_7 (points: [0])
```
>>> 100 <= triple_record_vert_est <= 200
True
```


Tests for q3_1 (points: [0, 0])
```
>>> 0 <= rsl_ns_slope < 1
True
```
```
>>> rsl_ns_intercept < 0
True
```


Tests for q3_2 (points: [])


Tests for q3_3 (points: [0])
```
>>> np.isclose(linear_model(0,0,0),0) 
True
```


Tests for q3_4 (points: [0])
```
>>> np.allclose(linear_fit(Table().with_column('x',np.array([1,2])).with_column('y',np.array([1,2]))),[1, 2])
True
```


Tests for q3_6 (points: [])


Tests for q3_7 (points: [0])
```
>>> np.isclose(linear_predictor(Table().with_column('x',np.array([1,2])).with_column('y',np.array([1,2])),3),3)
True
```


Tests for q3_8 (points: [])
