```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Lab 7: A/B Testing

Welcome to Data 111 Lab 7! This week's lab will focus on A/B Testing. We will begin with a discussion and exercise using data from Karuk and Yurok tribes. This discussion will include some heavy topics, so take time to ponder and share throughout this process. 

After our discussion, we will explore data from the ever-popular British television show, [*The Great British Bakeoff*](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off).


**Submission**: Once you’re finished, run all cells besides the last one, select File > Save Notebook, and then execute the final cell. Then submit the downloaded zip file, that includes your notebook..

First, set up the notebook by running the cell below.

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plots
plots.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', (FutureWarning, np.VisibleDeprecationWarning))

import d8error
```

# Discussion: Traditional Ecological Knowledge

Indigenous people are the people who have inhabited a land area from the earliest times, before colonization, and have a way of life that is *inextricably linked* to that land. Over many thousands of years and more, indegenous folks learn and pass down wisdom about how humans can best support the local ecosystem they are a part of and depend on. This wisdom is sometimes called "Traditional Ecological Knowledge" (TEK). Invasion by colonization across the world has caused many genocides of indigenous people and their lands, and has attempted to erase indigenous culture and TEK. Watch [this video](https://www.youtube.com/watch?v=T-azcPugmKQ) to see the impacts of colonialism on indigenous peoples here and discuss the following question:

**Discussion Question 1:** What are the ways that colonialism harmed/harms the indigenous people of our region?

"Western" science (like you are learning in this class) has been dominated by folks from colonialist cultures. It is a very young culture compared to indigenous culture, and it is perhaps not surprising that recent scientific studies are relearning the lessons of the much older indigenous wisdom. Watch [this video](https://www.youtube.com/watch?v=wlzDJQjjQDg) about indigenous cultural burning and wildfire management to answer the following question. 

**Discussion Question 2:** How do prescribed burns help with wildfire management?

In addition to protecting people and the ecosystems we are part of from dangerous wildfires, perscribed burns are needed for cultural resources that are essential to the Karuk and Yurok way of life. For this reason, Karuk and Yurok people (and many others indigenous peoples) refer to their prescribed fire as *cultural burns* that are central to subsistence and ceremonial practices. The term *cultural burning* distinguishes these fires from the fuel reduction-focused *prescribed burns* of public land agencies whose primary objective is to reduce fuel loads and moderate wildfire intensity. One of the ways cultural burns are central to Karuk and Yurok way of life, is that they provide more hazelnut stems that are needed for *basket weaving*, which we will learn more about in this activity.

**Discussion Question 3:** Why do Karuk and Yurok people prefer to call their prescribed burning practices "cultural burns"?

# Hazelnut Stems for Basketweaving
Baskets weaved from *hazelnut stems* have many uses in Karuk and Yurok culture, including (but not limited to) baby cradles, clothing, and food preparation. Here is a hopper basket used to pound acorns. The stems on the left side of the picture are unpeeled, whereas and peeled hazelnut basketry stems are shown on the right. This basket is composed of peeled hazelnut stems along with other materials. 

<img src="Basket.jpg" alt="drawing" width="300"/>

Skilled basketweavers gather the stems and braid the baskets. Here is a picture of Karuk basketweavers Janet Morehead (left) and Lillian Rentz (right), peeling and evaluating the quality of hazelnut stems they gathered.

<img src="Basketweavers.jpg" alt="drawing" width="400"/>

These basketweavers generously participated in a scientific study that compared the effect of different burn treatments on how hazelnut shrubs produce basket-worthy stems. Additional basketweavers involved in the study include Elizabeth Azzuz, Laverne Glaze, Chook-Chook Hillman, Lisa Hillman, Robert McConnell,Robert McConnell Jr., Kathy McCovey, Bertha Peters, Maggie Peters, Verna Reece, Margo Robbins, Brittany Souza, and Bill Tripp.

# Comparing Burn Treatments for Basket-Worthy Stem Production

We will look at some data from this study [Effects of understory fire management treatments on California Hazelnut, an ecocultural resource of the Karuk and Yurok Indians in the Pacific Northwest](https://www.sciencedirect.com/science/article/pii/S0378112719306826) published in *Forest Ecology and Managment* in 2019. 

Skilled basketweavers (and scientists trained by the skilled basketweavers) collected basket-worthy stems from many different hazelnut shrubs that each received a "treatment". One of the "treatments" was a "broadcast burn" which is the type of cultural burn that the Yurok and Karuk people prefer to use as part of their basketweaving traditions. In the "control," the shrub was left untouched. Stems were gathered before treatment and then again after a full growing season since the treatment. 

The Study Location (yellow star in the intersection of Karuk and Yurok Ancentral Lands) is shown on the map below.

<img src="AncestralTerritory.jpg" alt="drawing" width="700"/>

 Let's have a look at a table of the data...

```python
bb = Table.read_table('Basketweaving_and_Burning.csv')
```

```python
bb
```

**Discussion Question 4:** Were there more shrubs that experienced the "Broadcast" treatment or the "Control" treatment? How would you manipulate the table to figure this out?



```python
...
```

Let's plot the data to visualize it. We will plot the control treatment data using a slightly larger marker than the broadcast burn treatment data in case those data points end up falling on top of each other. We will also plot the data points with some transparency in case the data points within a treatment class fall on top of themselves.

```python
bb_control = bb.where('Treatment', are.equal_to('Control'))
bb_burn = bb.where('Treatment', are.equal_to('Broadcast'))
plots.plot([0,40],[0,40],'k-',linewidth=1)
plots.scatter(bb_control.column('PreBasketStems'),bb_control.column('PostBasketStems'),s=20,alpha=0.5,label='Control')
plots.scatter(bb_burn.column('PreBasketStems'),bb_burn.column('PostBasketStems'),s=10,alpha=0.5,label='Burn')
plots.xlabel('PreBasketStems')
plots.ylabel('PostBasketStems')
plots.legend(loc='upper right')
```

Use the plot above to answer the following discussion questions:

**Discussion Question 5:** What is plotted on the *x*-axis?

*Type your answer replacing this text*

**Discussion Question 6:** What is plotted on the *y*-axis?

*Type your answer replacing this text*

**Discussion Question 7:** What do the *blue* dots represent?

*Type your answer replacing this text*

**Discussion Question 8:** What do the *red* dots represent?

*Type your answer replacing this text*

**Discussion Question 9:** What does the *black line* represent?

*Type your answer replacing this text*

**Discussion Question 10:** Take a look at the two burned shrubs that started out with the most basket-worthy stems. (Red dots with the two largest x-values.) How many basket-worthy stems grew on those shrubs after the burn treatment? What do you think happened to those shrubs?

*Type your answer replacing this text*

**Discussion Question 11:** Are there any blue dots above the black line? How does this compare to the red dots? What does this mean?

```python

```

If you are curious about which datapoints fell on top of each other in the above plot, you can count them using the group function. First we look at the control data:

```python
bb_control.group(['PreBasketStems','PostBasketStems'])
```

And then we can look at the same thing in the broadcast burn data:

```python
bb_burn.group(['PreBasketStems','PostBasketStems'])
```

But perhaps it makes more sense to look at how many basket-worthy stems were on the shrubs after treatment compared to how many basket-worthy stems they started out with...

```python
bb = bb.with_column('BasketStemDiff',bb.column('PostBasketStems')-bb.column('PreBasketStems'))
stemdiff_and_treatment = bb.select('BasketStemDiff', 'Treatment')
```

```python
bb
```

Now we have added a new column showing the difference in basket-worthy stems before and after treatment. Now let's use that new column to make a histogram, where we separate out the data by treatment. We have use some transparency because the data from each treatment sometimes overlap.

```python
stemdiff_and_treatment.hist('BasketStemDiff', group='Treatment')
```

**Discussion Question 12:** What does a negative number on the x-axis mean?

*Type your answer replacing this text*

**Discussion Question 13:** Does it seem clear to you that one distribution is in general weighted farther to the right than the other? If so, which treatment seems best for creating more basket-worthy stems?

*Type your answer replacing this text*

# Chance Variation?
This raises the question of whether the difference in the distributions we're looking at is just due to chance or if the treatment has an effect of producing more basket-worthy stems.

We can try to answer this question through an A/B test. 

**Discussion Question 14:** State the *null hypothesis*, in your own words.

*Type your answer replacing this text*

**Discussion Question 15:** State the *alternative hypothesis*, in your own words.

*Type your answer replacing this text*

**Discussion Question 16:** What *test statistic* should we use to test the null hypothesis?

*Type your answer replacing this text*

# Test Statistic

Let's compute our test statistic. 

```python
means_table = bb.select(['Treatment','BasketStemDiff']).group('Treatment', np.average)
means_table
```

```python
means = means_table.column(1)
observed_difference = ...
observed_difference
```

**Discussion Question 17:** Based on how we have coded this, what values of our statistic are in favor of the alternative: positive or negative?

*Type your answer replacing this text*

Now let's create a function that can quickly calculate the test statistic for us. (It should do the same thing we just did above.)

```python
def difference_of_means(table, label_col, val_col):
    """Takes: name of table, column label of group-label variable, 
    and column label of numerical variable
    
    Returns: Difference of means of the two groups"""
    
    #table with the two relevant columns
    reduced = table.select(val_col, label_col)  
    
    # table containing group means
    means_table = reduced.group(label_col, np.average)
    # array of group means
    means = means_table.column(1)
    
    return ...
```

```python
difference_of_means(bb,'BasketStemDiff', 'Treatment')
```

How to we simulate data under the null hypothesis? Recall that shuffling up the labels (AKA random permutations) is a clever way of doing that. 

# Simulation Under Null Hypothesis

Create a new table Containing just the `BasketStemDiff` and the `Treatment` columns. 

```python
stemdiff_and_treatment = ...
```

Then add a column of to our table with the shuffled treatment labels

```python
shuffled_labels = ...
original_and_shuffled = stemdiff_and_treatment.with_column('Shuffled Label', ...)
original_and_shuffled
```

Now we can look at the difference of the group means...

```python
difference_of_means(original_and_shuffled, 'BasketStemDiff', 'Shuffled Label')
```

**Discussion Question 18:** How does this test statistic under the null hypothesis compare to the value of the test statistic that you found in the real observations?

*Type your answer replacing this text*

But we only shuffled the labels once. Could a different shuffle have resulted in a larger difference between the group averages? To get a sense of the variability, we must simulate the difference many times. Let's shuffle them a bunch of times to get an estimate of the *distribution* of the test statistic under the null for each shuffling that we perform.

# Permutation Test

We will start by defining a function that simulates one value of the test statistic under the null hypothesis. This is just a matter of collecting the code that we wrote above.

```python
def one_simulated_difference(table, label_col, val_col):
    """Takes: name of table, column label of group-label variable, 
    and column label of numerical variable
    
    Returns: Difference of means of the two groups after shuffling labels"""
    
    # array of shuffled labels
    shuffled_labels = ...
    
    # table of numerical variable and shuffled labels
    shuffled_table = ...
    
    return ...   
```

Run the cell below multiple times to see how the test statistic varies under the null hypothesis:

```python
one_simulated_difference(stemdiff_and_treatment, 'BasketStemDiff', 'Treatment')
```

```python
differences = make_array()

for i in np.arange(2500):
    new_difference = one_simulated_difference(stemdiff_and_treatment, 'BasketStemDiff', 'Treatment')
    differences = np.append(differences, new_difference)
```

```python
Table().with_column('Difference Between Group Means', differences).hist()
print('Observed Difference:', observed_difference)
plots.title('Prediction Under the Null Hypothesis');
plots.scatter(observed_difference, 0, color='red', s=40, zorder=3)
```

**Discussion Question 19:** How does the observed test statistic compare to the distribution of the statistic under the null hypothesis? 

*Type your answer replacing this text*

**Discussion Question 20:** Which hypothesis do the data best support?

*Type your answer replacing this text*

**Discussion Question 21:** In your own words, describe the effect of cultural burning on basket-worthy stem growth, as supported by the data.

*Type your answer replacing this text*

# How things are now...

Despite the now widespread acknowledgement of the wisdom of cultural burning, these practices are restricted due USDA Forest Service funding and staff reductions, threatened and endangered wildlife being protected in certain seasons, burn restrictions during major wildfire events, and air quality regulations that constrain available burn days. 

Basketweavers also say the departure of supportive USDA Forest Service managers created major set-backs, because they had to re-establish lines of communication with new staff and inform them about their basketry materials and fire treatment
needs. So even with the federal governments increased fuel reduction treatments throughout this region basketweavers and Tribal members often were not informed about the schedules or locations of understory mechanical treatments and broadcast burns, and would miss opportunities to gather hazelnut stems. 

Efforts within the USDA Forest Service known as the ‘Shared Stewardship’ initiative seek to address several of these constraints.

In the meantime basketweavers have been coming up with their own creative solutions to stimulate basket-worthy stem growth on hazelnut shrubs. The figure below shows different types of treatments now being used: (A) manual cutting of the shrub; (B) small pile burn on an indivual shrub; (C) propone torching of a shrub (D) the traditional broadcast burn.
<img src="Treatments.jpg" alt="drawing" width="400"/>

All of these treatments increase basket-worthy stem growth, where the ones that use fire are the most effective. You can read more about the details in [the paper](https://www.sciencedirect.com/science/article/pii/S0378112719306826) if you like.


At home, watch [this video](https://www.youtube.com/watch?v=QQfrQyb4yw4) about the *Land Back* movement. If you want to read more, here is a [news article](https://www.npr.org/2020/08/24/899422710/to-manage-wildfire-california-looks-to-what-tribes-have-known-all-along) that also covered this topic. 

## 1. A/B Testing

A/B testing is a form of hypothesis testing that allows you to make comparisons between two distributions. We may also refer to an A/B test as a permutation test.

You'll almost never be explicitly asked to perform an A/B test. Make sure you can identify situations where the test is appropriate and know how to correctly implement each step. Oftentimes, we use an A/B test to determine whether or not two samples came from the same underlying distribution.

**Question 1.1.** The following statements are the steps of an A/B hypothesis test presented in a *random order*:

1. Choose a test statistic (typically the difference in means between two categories)

2. Shuffle the labels of the original sample, find your simulated test statistic, and repeat many times

3. Find the value of the observed test statistic

4. Calculate the p-value based off your observed and simulated test statistics

5. Define a null and alternate model

6. Use the p-value and p-value cutoff to draw a conclusion about the null hypothesis

Assign `ab_test_order` to an array of integers that contains the correct order of an A/B test, where the first item of the array is the first step of an A/B test and the last item of the array is the last step of an A/B test.


```python
ab_test_order = ...
```

```python
grader.check("q1_1")
```

**Question 1.2.** If the null hypothesis of an A/B test is correct, should the order of labels affect the differences in means between each group? Why do we shuffle labels in an A/B test? If you are in a lab section, confirm your answer with a neighbor or staff member before moving on. 


_Type your answer here, replacing this text._

## 2. The Great British Bake Off

>"The Great British Bake Off (often abbreviated to Bake Off or GBBO) is a British television baking competition, produced by Love Productions, in which a group of amateur bakers compete against each other in a series of rounds, attempting to impress a group of judges with their baking skills" [Wikipedia](https://en.wikipedia.org/wiki/The_Great_British_Bake_Off)

For every week of the competition, the judges assign one contestant the title "Star Baker". Ultimately, one winner is crowned every season. Using this information, we would like to investigate how winning Star Baker awards affects the odds of winning a season of the show.

**Question 2.1.** We want to know whether winning more Star Baker awards *causes* a change in likelihood of winning the season.  Why is it not sufficient to compare star baker rates for winners and losers?


_Type your answer here, replacing this text._

### Running an Experiment

We are going to run the following hypothesis test to determine the association between winning and number of Star Baker awards. The population we are examining is every contestant from seasons 2 through 11 of GBBO. We are going to use the following null and alternative hypotheses:

**Null hypothesis:** The distribution of Star Baker awards between contestants who won their season and contestants who did not win their season is the same.

**Alternative hypothesis:** Contestants who win their season of the show will win more Star Baker awards on average.

Our alternative hypothesis is related to our suspicion that contestants who win more Star Baker awards are more skilled, so they are more likely to win the season.

**Question 2.2.** Should we use an A/B test to test these hypotheses? If yes, what is our "A" group and what is our "B" group?


_Type your answer here, replacing this text._

Check your answers with your neighbors or a staff member before you move on to the next section.

The `bakers` table below describes the number of star baker awards each contest won and whether or not they won their season (`1` if they won, `0` if they did not win). The data was manually aggregated from Wikipedia for seasons 2-11 of the show. We randomized the order of rows as to not spoil the outcome of the show.

```python
bakers = Table.read_table("star_bakers.csv")
bakers.show(3)
```

**Question 2.3.** Create a new table called `means` that contains the mean number of star baker awards for bakers who did not win (`won==0`) and bakers that did win (`won==1`). The table should have the column names `won` and `star baker awards mean`.

```python
means = ...
means
```

```python
grader.check("q2_3")
```

**Question 2.4.** Visualize the distribution of Star Baker awards for winners and non-winners. You should use the bins we provided.

Hint: You will want to use the group argument of `tbl.hist`. In order to produce several overlayed histograms based on unique values in a given column, we can do something like `tbl.hist(..., group=<col_name>, bins=...)`!


```python
useful_bins = np.arange(0, 7)
...
```

**Question 2.5.** We want to figure out if there is a difference between the distribution of Star Baker awards between winners and non winners. 

What should the test statistic be? Which values of this test statistic support the null, and which values support the alternative?

If you are in lab, confirm your answer with a neighbor or staff member before moving on.

Hint: You should think about what measures we use to describe a distribution. 


_Type your answer here, replacing this text._

**Question 2.6.** Set `observed_difference` to the observed test statistic using the `means` table. 


```python
observed_difference = ...
observed_difference
```

```python
grader.check("q2_6")
```

**Question 2.7.** Given a table like `bakers`, a label column `label_col`, and a values column `val_col`, write a function that calculates the appropriate test statistic.

*Hint:* Make sure that you are taking the directionality of our alternative hypothesis into account.


```python
def find_test_stat(tbl, label_col, val_col):
    ...

find_test_stat(bakers, "won", "star baker awards")
```

```python
grader.check("q2_7")
```

When we run a simulation for A/B testing, we resample by **shuffling the labels** of the original sample. If the null hypothesis is true and the star baker award distributions are the same, we expect that the difference in mean star baker awards will be not change when `"won"` labels are changed.

**Question 2.8.** Write a function `simulate_and_test_statistic` to compute one trial of our A/B test. Your function should run a simulation and return a test statistic.


```python
def simulate_and_test_statistic(tbl, labels_col, values_col):
    ...

simulate_and_test_statistic(bakers, "won", "star baker awards")
```

```python
grader.check("q2_8")
```

**Question 2.9.** Simulate 5000 trials of our A/B test and store the test statistics in an array called `differences`.


```python
# This cell might take a couple seconds to run
differences = make_array()

...
                                                 
differences
```

```python
grader.check("q2_9")
```

Run the cell below to view a histogram of your simulated test statistics plotted with your observed test statistic.

```python
Table().with_column('Difference Between Group Means', differences).hist(bins=20)
plots.scatter(observed_difference, 0, color='red', s=30, zorder=2)
plots.ylim(-0.1, 1.35);
```

**Question 2.10.** Find the p-value for your test and assign it to `empirical_p`.


```python
empirical_p = ...
empirical_p
```

```python
grader.check("q2_10")
```

**Question 2.11.** Using a 5% P-value cutoff, draw a conclusion about the null and alternative hypotheses. Describe your findings using simple, non-technical language. What does your analysis tell you about the association between star baker awards and winning? What can you claim about causation from your statistical analysis? Call your GSI or AI over and confirm your answer. 


_Type your answer here, replacing this text._

```python
grader.check("q2_10")
```

## 3. Submission

Gus wanted to say congrats on finishing Lab 7!! 

<img src="gus_gus.jpg" alt="drawing" width="300"/>


**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Then submit the zip file to the corresponding assignment in Canvas. 

**It is your responsibility to make sure your work is saved before running the last cell.**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: None)
```
>>> len(ab_test_order) == 6
True
```
```
>>> np.allclose(np.array(ab_test_order)[:3] % 2 == 1, True) # Check the first three elements in your array.
True
```
```
>>> np.allclose(np.array(ab_test_order)[3:] % 2 == 0, True) # Check the last three elements in your array.
True
```
```
>>> 
>>> # This imports a hashing library for the autograder.
>>> import hashlib
>>> def get_hash(num):
...     """Helper function for assessing correctness."""
...     return hashlib.md5(str(num).encode()).hexdigest()
>>> 
>>> get_hash(np.array(ab_test_order).astype(int))
'a7196ed0f271c873d9750cb92422d911'
```


Tests for q2_10 (points: None)
```
>>> 0 <= empirical_p < 0.05
True
```


Tests for q2_3 (points: None)
```
>>> means.num_rows
2
```
```
>>> np.round(min(means.column("star baker awards mean")), 2) == 0.65
True
```
```
>>> np.round(max(means.column("star baker awards mean")), 2) == 1.5
True
```


Tests for q2_6 (points: None)
```
>>> isinstance(observed_difference, float)
True
```
```
>>> float(round(observed_difference, 3))
0.848
```


Tests for q2_7 (points: None)
```
>>> np.isclose(round(find_test_stat(bakers, "won", "star baker awards"), 3) - 0.848, 0)
True
```


Tests for q2_8 (points: None)
```
>>> test_stat = round(simulate_and_test_statistic(bakers, "won", "star baker awards"), 3)
>>> -2 < test_stat < 2
True
```
```
>>> np.random.seed(1)
>>> test_stat2 = simulate_and_test_statistic(bakers, "won", "star baker awards")
>>> np.round(test_stat2, 3) == -0.023 or np.round(test_stat2, 3) == -0.132
True
```


Tests for q2_9 (points: None)
```
>>> len(differences)
5000
```
```
>>> abs(np.average(differences)) < 0.05 # On average, your test statistic should be close to 0
True
```
```
>>> all(differences == differences.item(0)) == False # Make sure all of the test statistics are different
True
```
